# MISA
Alohan'ny mamerina dia avereno atao Run ny notebook iray manontolo. Ny fanaovana azy dia redémarrena mihitsy ny kernel aloha (jereo menubar, safidio **Kernel$\rightarrow$Restart Kernel and Run All Cells**).

Izay misy hoe `YOUR CODE HERE` na `YOUR ANSWER HERE` ihany no fenoina. Afaka manampy cells vaovao raha ilaina. Aza adino ny mameno references eo ambany raha ilaina.

## References
Eto ilay references rehetra no apetraka

---

```python
from random import randrange
import numpy as np
from sklearn.metrics import mean_squared_error, log_loss
from sklearn.datasets import load_breast_cancer, load_diabetes


def grad_check_sparse(f, x, analytic_grad, num_checks=10, h=1e-5, error=1e-9):
    """
    sample a few random elements and only return numerical
    in this dimensions
    """

    for i in range(num_checks):
        ix = tuple([randrange(m) for m in x.shape])

        oldval = x[ix]
        x[ix] = oldval + h  # increment by h
        fxph = f(x)  # evaluate f(x + h)
        x[ix] = oldval - h  # increment by h
        fxmh = f(x)  # evaluate f(x - h)
        x[ix] = oldval  # reset

        grad_numerical = (fxph - fxmh) / (2 * h)
        grad_analytic = analytic_grad[ix]
        rel_error = abs(grad_numerical - grad_analytic) / (
            abs(grad_numerical) + abs(grad_analytic)
        )
        print(
            "numerical: %f analytic: %f, relative error: %e"
            % (grad_numerical, grad_analytic, rel_error)
        )
        assert rel_error < error

def rel_error(x, y):
    """ returns relative error """
    return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))
```

# Linear regression

```python
data = load_diabetes()
X_train1, y_train1 = data.data, data.target
w1 = np.random.randn(X_train1.shape[1]) * 0.0001
b1 = np.random.randn(1) * 0.0001
```

<!-- grade_id: cell-5ffdf99ad7cdfd69, points: 0, grade: false, solution: true -->
```python
def mse_loss_naive(w, b, X, y, alpha=0):
    """
    MSE loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    # YOUR CODE HERE
   
    for i in range(len(y)):
        loss += ((1 / len(y)) * ((w.T).dot(X[i]) + b - y[i]) ** 2)
        db += 2 * (1 / len(y)) * ((w.T).dot(X[i]) + b - y[i])
    loss+=alpha*(w.T).dot(w)+alpha*b**2
    db+=2*alpha*b
    for i in range (len(w)):
        for j in range (len(y)):
            dw[i] += 2 * (1 / len(y)) * (X[j][i]) * ((w.T).dot(X[j]) + b - y[j])
        dw[i]+=2*(alpha)*w[i]
    return loss, dw, np.array(db).reshape(1,)
```

## Naive Linear regression loss

<!-- grade_id: cell-079e42b153b67a60, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  5.005050726161976e-16
Gradient check w
numerical: -3.153316 analytic: -3.153317, relative error: 1.631872e-07
numerical: -1.275043 analytic: -1.275044, relative error: 5.365633e-07
numerical: -3.153316 analytic: -3.153317, relative error: 1.631872e-07
numerical: -1.275043 analytic: -1.275044, relative error: 5.365633e-07
numerical: -4.145418 analytic: -4.145418, relative error: 5.119551e-08
numerical: -1.553189 analytic: -1.553189, relative error: 2.987241e-08
numerical: -1.275043 analytic: -1.275044, relative error: 5.365633e-07
numerical: -2.801914 analytic: -2.801913, relative error: 1.249387e-07
numerical: -2.801914 analytic: -2.801913, relative error: 1.249387e-07
numerical: -2.801914 analytic: -2.801913, relative error: 1.249387e-07
numerical: -0.315455 analytic: -0.315453, relative error: 2.038005e-06
numerical: -1.376395 analytic: -1.376394, relative error: 3.231708e-07
numerical: -1.376395 analytic: -1.376394, relative error: 3.231708e-07
numerical: -1.275043 ana
```

## Naive Ridge regression loss

<!-- grade_id: cell-14a64cb1cfe70b98, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -0.315255 analytic: -0.315254, relative error: 1.959728e-06
numerical: -3.153142 analytic: -3.153143, relative error: 1.605151e-07
numerical: -2.802017 analytic: -2.802016, relative error: 1.184452e-07
numerical: 2.891921 analytic: 2.891922, relative error: 1.166039e-07
numerical: 2.891921 analytic: 2.891922, relative error: 1.166039e-07
numerical: -2.802017 analytic: -2.802016, relative error: 1.184452e-07
numerical: -1.275430 analytic: -1.275432, relative error: 5.472000e-07
numerical: -4.296376 analytic: -4.296377, relative error: 3.835617e-08
numerical: -3.234021 analytic: -3.234020, relative error: 1.687704e-07
numerical: -2.802017 analytic: -2.802016, relative error: 1.184452e-07
numerical: -0.315255 analytic: -0.315254, relative error: 1.959728e-06
numerical: -1.376237 analytic: -1.376236, relative error: 3.622452e-07
numerical: -2.802017 analytic: -2.802016, relative error: 1.184452e-07
numerical: -1.376237 analytic: -1.376236, relative error: 3.6224
```

<!-- grade_id: cell-1528a28f467d90c7, points: 0, grade: false, solution: true -->
```python
def mse_loss_vectorized(w, b, X, y, alpha=0):
    """
    MSE loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    loss = 0.0
    dw = np.zeros_like(w)
    ones =np.ones((X.shape[0],1))
    X = np.hstack((ones,X))
    w = np.concatenate((np.array(b).reshape(1,),w),axis=0)
    
    # YOUR CODE HERE
    v=np.dot(X,w)-y
    loss = np.dot(v.T,v)/len(y)
    deriv = 2/len(y)*(X.T).dot((X.dot(w)-y))
    deriv += 2*(alpha)*w
    dw = deriv[1:]
    db = deriv[0]
    
    loss+=alpha*(w.T).dot(w)+alpha*b**2
    
    
    return loss, dw, np.array(db).reshape(1,)
```

## Vectorised Linear regression loss

<!-- grade_id: cell-41637ca21c8f938d, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  0.0
Gradient check w
numerical: -2.801913 analytic: -2.801913, relative error: 4.900373e-09
numerical: -2.801913 analytic: -2.801913, relative error: 4.900373e-09
numerical: -3.153317 analytic: -3.153317, relative error: 9.867661e-09
numerical: -4.145418 analytic: -4.145418, relative error: 2.925575e-08
numerical: -4.145418 analytic: -4.145418, relative error: 2.925575e-08
numerical: -3.153317 analytic: -3.153317, relative error: 9.867661e-09
numerical: -1.376394 analytic: -1.376394, relative error: 7.219510e-09
numerical: -4.296088 analytic: -4.296088, relative error: 6.365543e-09
numerical: -4.145418 analytic: -4.145418, relative error: 2.925575e-08
numerical: -3.153317 analytic: -3.153317, relative error: 9.867661e-09
numerical: -4.145418 analytic: -4.145418, relative error: 2.925575e-08
numerical: -4.145418 analytic: -4.145418, relative error: 2.925575e-08
numerical: -1.275044 analytic: -1.275044, relative error: 3.408044e-08
numerical: -1.376394 analytic: -1.376394, 
```

## Vectorized ridge regression loss

<!-- grade_id: cell-bb0ff99c19a85e0f, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -1.376236 analytic: -1.376236, relative error: 3.181695e-08
numerical: -2.802016 analytic: -2.802016, relative error: 1.138912e-08
numerical: -3.234020 analytic: -3.234020, relative error: 2.815681e-08
numerical: -1.275432 analytic: -1.275432, relative error: 2.327048e-08
numerical: -3.153143 analytic: -3.153143, relative error: 1.254930e-08
numerical: 2.891922 analytic: 2.891922, relative error: 9.194056e-09
numerical: -3.234020 analytic: -3.234020, relative error: 2.815681e-08
numerical: -3.234020 analytic: -3.234020, relative error: 2.815681e-08
numerical: -1.275432 analytic: -1.275432, relative error: 2.327048e-08
numerical: -3.234020 analytic: -3.234020, relative error: 2.815681e-08
numerical: -0.315254 analytic: -0.315254, relative error: 5.973995e-08
numerical: -2.802016 analytic: -2.802016, relative error: 1.138912e-08
numerical: -1.376236 analytic: -1.376236, relative error: 3.181695e-08
numerical: -2.802016 analytic: -2.802016, relative error: 1.13
```

# Logistic regression

```python
def sigmoid(z):
    return 1 / (1 + np.exp(-z))

data = load_breast_cancer()
X_train2, y_train2 = data.data, data.target
w2 = np.random.randn(X_train2.shape[1]) * 0.0001
b2 = np.random.randn(1) * 0.0001
```

# Naive

<!-- grade_id: cell-6ec156568b0c6e29, points: 0, grade: false, solution: true -->
```python
def log_loss_naive(w, b, X, y, alpha=0):
    """
    log loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    # YOUR CODE HERE
     
    
    for i in range(len(y)):
        z = (w.T).dot(X[i])+b
        A= sigmoid(z)
        loss += (-y[i] * np.log(A) - (1 - y[i]) * np.log(1 - A))
        
        db += (A - y[i])
        for k in range(len(w)):
            dw[k]+=(A-y[i])*X[i][k]
    loss =loss/len(y)
    db = db/len(y)
    dw= dw/len(y)
    loss+=alpha*(w.T).dot(w)+alpha*b**2
    db+=2*alpha*b
    dw+=2*(alpha)*w
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-d3c078eb2449ee61, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  7.70240915883959e-17
Gradient check w
numerical: 0.742211 analytic: 0.742211, relative error: 2.507711e-09
numerical: 10.191149 analytic: 10.191150, relative error: 4.853932e-08
numerical: -0.000187 analytic: -0.000187, relative error: 3.445378e-08
numerical: 10.191149 analytic: 10.191150, relative error: 4.853932e-08
numerical: 5.698381 analytic: 5.698383, relative error: 9.780162e-08
numerical: 82.211410 analytic: 82.211914, relative error: 3.063762e-06
numerical: 0.052169 analytic: 0.052169, relative error: 9.053593e-11
numerical: 0.000737 analytic: 0.000737, relative error: 6.411997e-09
numerical: 0.027995 analytic: 0.027995, relative error: 1.129124e-09
numerical: 5.698381 analytic: 5.698383, relative error: 9.780162e-08
numerical: 153.752250 analytic: 153.753996, relative error: 5.678853e-06
numerical: -0.003445 analytic: -0.003445, relative error: 1.237264e-09
numerical: -0.004819 analytic: -0.004819, relative error: 8.220889e-09
numerical: -0.095170 analytic: -0.0
```

# Naive with regulariztion

<!-- grade_id: cell-b91805308ad91ec5, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: 0.011873 analytic: 0.011873, relative error: 4.699986e-10
numerical: -0.009270 analytic: -0.009270, relative error: 2.489930e-10
numerical: 0.000609 analytic: 0.000609, relative error: 1.909775e-08
numerical: 82.211540 analytic: 82.212044, relative error: 3.063757e-06
numerical: 2.545900 analytic: 2.545901, relative error: 1.185248e-07
numerical: 0.022005 analytic: 0.022005, relative error: 1.028317e-09
numerical: 0.288017 analytic: 0.288017, relative error: 3.719337e-09
numerical: 0.001295 analytic: 0.001295, relative error: 1.175466e-08
numerical: 0.288017 analytic: 0.288017, relative error: 3.719337e-09
numerical: 0.018397 analytic: 0.018397, relative error: 1.410402e-10
numerical: -0.000743 analytic: -0.000743, relative error: 7.898853e-09
numerical: -0.095314 analytic: -0.095314, relative error: 1.645748e-10
numerical: -0.003481 analytic: -0.003481, relative error: 1.341357e-09
numerical: 153.752628 analytic: 153.754374, relative error: 5.678840e-06
num
```

# Vectorized

<!-- grade_id: cell-a96e9a6d51919ffc, points: 0, grade: false, solution: true -->
```python
def log_loss_vectorized(w, b,X, y, alpha=0):
    """
    log loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    """
    loss = 0.0
    dw = np.zeros_like(w)
    
    # YOUR CODE HERE
    Z = X.dot(w)+ b 
    A= sigmoid(Z)
    loss =1/len(y) * np.sum(-y * np.log(A)-(1-y)*np.log(1-A))
    dw=1/len(y)*np.dot(X.T,A-y)
    db = 1/len(y)*np.sum(A-y)
    loss+=alpha*(w.T).dot(w)+alpha*b**2
    db+=2*alpha*b
    dw+=2*(alpha)*w
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-ca14e49ada130789, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  0.0
Gradient check w
numerical: 0.011895 analytic: 0.011895, relative error: 2.480508e-11
numerical: -0.095170 analytic: -0.095170, relative error: 2.664265e-11
numerical: 82.211410 analytic: 82.211914, relative error: 3.063762e-06
numerical: -0.000187 analytic: -0.000187, relative error: 9.977569e-09
numerical: 0.018160 analytic: 0.018160, relative error: 8.233621e-11
numerical: -0.004819 analytic: -0.004819, relative error: 4.184906e-10
numerical: 82.211410 analytic: 82.211914, relative error: 3.063762e-06
numerical: 153.752250 analytic: 153.753996, relative error: 5.678853e-06
numerical: 0.288417 analytic: 0.288417, relative error: 3.644901e-09
numerical: 82.211410 analytic: 82.211914, relative error: 3.063762e-06
numerical: -0.004716 analytic: -0.004716, relative error: 3.106722e-10
numerical: 0.742211 analytic: 0.742211, relative error: 2.507712e-09
numerical: 0.011895 analytic: 0.011895, relative error: 2.480508e-11
numerical: 0.021882 analytic: 0.021882, relative e
```

# Vectorized with regularization

<!-- grade_id: cell-bce082fccf8b7057, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -0.003481 analytic: -0.003481, relative error: 2.531082e-10
numerical: 0.049958 analytic: 0.049958, relative error: 3.277218e-13
numerical: -0.000096 analytic: -0.000096, relative error: 1.795359e-08
numerical: 0.375425 analytic: 0.375425, relative error: 1.152036e-10
numerical: 0.018397 analytic: 0.018397, relative error: 1.607025e-10
numerical: 153.752628 analytic: 153.754374, relative error: 5.678840e-06
numerical: -0.095314 analytic: -0.095314, relative error: 1.897485e-11
numerical: 82.211540 analytic: 82.212044, relative error: 3.063757e-06
numerical: -0.001312 analytic: -0.001312, relative error: 5.813581e-10
numerical: 0.288017 analytic: 0.288017, relative error: 3.651879e-09
numerical: 0.000609 analytic: 0.000609, relative error: 3.674196e-09
numerical: 0.288017 analytic: 0.288017, relative error: 3.651879e-09
numerical: -0.541181 analytic: -0.541181, relative error: 3.648502e-09
numerical: -0.004976 analytic: -0.004976, relative error: 3.671134e-10
```

# Gradient descent for Linear models

<!-- grade_id: cell-485e52c0efd4f4a9, points: 0, grade: false, solution: true -->
```python
class LinearModel():
    def __init__(self):
        self.w = None
        self.b = None

    def train(self, X, y, learning_rate=1e-3, alpha=0, num_iters=100, batch_size=200, verbose=False):
        N, d = X.shape
        
        if self.w is None: # Initialization
            self.w = 0.001 * np.random.randn(d)
            self.b = 0.0

        # Run stochastic gradient descent to optimize w
        
        loss_history = []
        for it in range(num_iters):
            X_batch = None
            y_batch = None
                                                               
            # Sample batch_size elements in X_batch and y_batch
            # X_batch shape is  (batch_size, d) and y_batch shape is (batch_size,)                                                                                          
            # Hint: Use np.random.choice to generate indices
            # YOUR CODE HERE
            index = np.random.choice(N,batch_size,replace=False)
            X_batch = np.array([X[i] for i in index])
            y_batch = np.array([y[i] for i in index])
            
            
            # evaluate loss and gradient
            loss, dw, db = self.loss(X_batch, y_batch, alpha)
            loss_history.append(loss)

            # perform parameter update                                                                
            # Update the weights w and bias b using the gradient and the learning rate.          
            # YOUR CODE HERE
            self.w= self.w- learning_rate * dw 
            self.b = self.b- learning_rate * db 
                    
            if verbose and it % 10000 == 0:
                print("iteration %d / %d: loss %f" % (it, num_iters, loss))
                
        return loss_history

    def predict(self, X):
        pass

    def loss(self, X_batch, y_batch, reg):
        pass

class LinearRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return mse_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        # YOUR CODE HERE
        y = X.dot(self.w) + self.b
        return y

        
class LogisticRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return log_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        """ Return prediction labels vector of 0 or 1 """
        # YOUR CODE HERE
        y = X.dot(self.w) + self.b
        r = np.zeros_like(y)
        for i in range(len(y)):
            r[i] = 1 if y[i]>0 else 0
        return r
        
```

## Linear regression with gradient descent

<!-- grade_id: cell-92f36a3b387a4277, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LinearRegression

sk_model = LinearRegression(fit_intercept=True)
sk_model.fit(X_train1, y_train1)
sk_pred = sk_model.predict(X_train1)
sk_mse = mean_squared_error(sk_pred, y_train1)

model = LinearRegressor()
model.train(X_train1, y_train1, num_iters=75000, batch_size=64, learning_rate=1e-2, verbose=True)
pred = model.predict(X_train1)
mse = mean_squared_error(pred, y_train1)

print("MSE scikit-learn:", sk_mse)
print("MSE gradient descent model :", mse)
assert mse - sk_mse < 100
```

Output:
```
iteration 0 / 75000: loss 27084.531224
iteration 10000 / 75000: loss 3066.402594
iteration 20000 / 75000: loss 2812.985126
iteration 30000 / 75000: loss 2943.337115
iteration 40000 / 75000: loss 2548.353319
iteration 50000 / 75000: loss 3001.407949
iteration 60000 / 75000: loss 2525.551435
iteration 70000 / 75000: loss 4153.144433
MSE scikit-learn: 2859.6963475867506
MSE gradient descent model : 2884.3536081730335
```

## Logistc regression with gradient descent

<!-- grade_id: cell-925a2ddb5c1ba7ff, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LogisticRegression
from sklearn.preprocessing import StandardScaler

scaler = StandardScaler()
X_train2 = scaler.fit_transform(X_train2)

sk_model = LogisticRegression(fit_intercept=True)
sk_model.fit(X_train2, y_train2)
sk_pred = sk_model.predict(X_train2)
sk_log_loss = log_loss(sk_pred, y_train2)

model = LogisticRegressor()
model.train(X_train2, y_train2, num_iters=75000, batch_size=64, learning_rate=1e-3, verbose=True)
pred = model.predict(X_train2)
model_log_loss = log_loss(pred, y_train2)

print("Log-loss scikit-learn:", sk_log_loss)
print("Log-loss gradiet descent model :", model_log_loss)
print("Error :", rel_error(sk_log_loss, model_log_loss))
assert rel_error(sk_log_loss, model_log_loss) < 1e-7
```

Output:
```
iteration 0 / 75000: loss 0.692680
iteration 10000 / 75000: loss 0.094124
iteration 20000 / 75000: loss 0.052587
iteration 30000 / 75000: loss 0.102940
iteration 40000 / 75000: loss 0.160897
iteration 50000 / 75000: loss 0.083378
iteration 60000 / 75000: loss 0.120880
iteration 70000 / 75000: loss 0.023333
Log-loss scikit-learn: 0.44341928598210933
Log-loss gradiet descent model : 0.44341928598210933
Error : 0.0
```

```python

```